```python
import numpy as np
import keras
from keras import layers
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[1], line 2
      1 import numpy as np
----> 2 import keras
      3 from keras import layers

ModuleNotFoundError: No module named 'keras'
```

<!-- grade_id: cell-0f6dc784a511426d, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

# Model / data parameters
num_classes = 10
input_shape = (28, 28, 1)

# Load the data and split it between train and test sets
(x_train, y_train), (x_test, y_test) = keras.datasets.mnist.load_data()

# Scale images to the [0, 1] range
x_train = x_train.astype("float32") / 255
x_test = x_test.astype("float32") / 255

# Make sure images have shape (28, 28, 1)
x_train = np.expand_dims(x_train, -1)
x_test = np.expand_dims(x_test, -1)

### END SOLUTION

```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[7], line 8
      5 input_shape = (28, 28, 1)
      7 # Load the data and split it between train and test sets
----> 8 (x_train, y_train), (x_test, y_test) = keras.datasets.mnist.load_data()
     10 # Scale images to the [0, 1] range
     11 x_train = x_train.astype("float32") / 255

NameError: name 'keras' is not defined
```

```python
# convert class vectors to binary class matrices
y_train = keras.utils.to_categorical(y_train, num_classes)
y_test = keras.utils.to_categorical(y_test, num_classes)

```

```python
model = keras.Sequential(
    [
        keras.Input(shape=input_shape),
        layers.Conv2D(32, kernel_size=(3, 3), activation="relu"),
        layers.MaxPooling2D(pool_size=(2, 2)),
        layers.Conv2D(64, kernel_size=(3, 3), activation="relu"),
        layers.MaxPooling2D(pool_size=(2, 2)),
        layers.Flatten(),
        layers.Dropout(0.5),
        layers.Dense(num_classes, activation="softmax"),
    ]
)

model.summary()
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[3], line 1
----> 1 model = keras.Sequential(
      2     [
      3         keras.Input(shape=input_shape),
      4         layers.Conv2D(32, kernel_size=(3, 3), activation="relu"),
      5         layers.MaxPooling2D(pool_size=(2, 2)),
      6         layers.Conv2D(64, kernel_size=(3, 3), activation="relu"),
      7         layers.MaxPooling2D(pool_size=(2, 2)),
      8         layers.Flatten(),
      9         layers.Dropout(0.5),
     10         layers.Dense(num_classes, activation="softmax"),
     11     ]
     12 )
     14 model.summary()

NameError: name 'keras' is not defined
```

<!-- grade_id: cell-915c43253eb81f43, points: 10, grade: true, solution: false -->
```python
### HIDDEN TESTS

# Check for MaxPooling2D layers
found_maxpooling2d = 0
for layer in model.layers:
    if isinstance(layer, layers.MaxPooling2D):
        found_maxpooling2d += 1
        assert layer.pool_size == (2, 2), f"Unexpected pool_size: {layer.pool_size}"

# Assert that there are exactly 2 MaxPooling2D layers
assert found_maxpooling2d == 2, f"Expected 2 MaxPooling2D layers, but found {found_maxpooling2d}"

print("All assertions passed.")

### HIDDEND TESTS
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[8], line 5
      1 ### HIDDEN TESTS
      2 
      3 # Check for MaxPooling2D layers
      4 found_maxpooling2d = 0
----> 5 for layer in model.layers:
      6     if isinstance(layer, layers.MaxPooling2D):
      7         found_maxpooling2d += 1

NameError: name 'model' is not defined
```

```python

```

```python
batch_size = 128
epochs = 1

model.compile(loss="categorical_crossentropy", optimizer="adam", metrics=["accuracy"])

model.fit(x_train, y_train, batch_size=batch_size, epochs=epochs, validation_split=0.1)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[4], line 4
      1 batch_size = 128
      2 epochs = 1
----> 4 model.compile(loss="categorical_crossentropy", optimizer="adam", metrics=["accuracy"])
      6 model.fit(x_train, y_train, batch_size=batch_size, epochs=epochs, validation_split=0.1)

NameError: name 'model' is not defined
```

```python
score = model.evaluate(x_test, y_test, verbose=0)
print("Test loss:", score[0])
print("Test accuracy:", score[1])
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[5], line 1
----> 1 score = model.evaluate(x_test, y_test, verbose=0)
      2 print("Test loss:", score[0])
      3 print("Test accuracy:", score[1])

NameError: name 'model' is not defined
```

```python

```