#### Write a function that returns sum of two numbers, Make sure it handles exception case

<!-- grade_id: cell-5beb3b8b93b61359, points: 0, grade: false, solution: true -->
```python
def computeSum(a, b):
    """Compute the sum of two numbers.
    
    """
    
    ### BEGIN SOLUTION
    if a == 0 or b == 0:
        raise ValueError("a & b both must be greater than or less than to 0")
        
    return a+b
    ### END SOLUTION
  
```

```python
computeSum(88, 11)
```

Output:
```
99
```

<!-- grade_id: cell-b4915a4279b3dcca, points: 10, grade: true, solution: false -->
```python
assert computeSum(14, 5) == 19
assert computeSum(2, 47) == 49
```

<!-- grade_id: cell-87031aafaa9c3e48, points: 10, grade: true, solution: false -->
```python
"""Check that computeSum returns the correct output for serveral inputs"""
### BEGIN HIDDEN TESTS
assert computeSum(4, 5) == 9
assert computeSum(2, 87) == 89
### END HIDDEN TESTS
```

```python
"""Check that computeSum raises an error for invalid inputs"""
try:
    computeSum(4,0)
except ValueError:
    pass
else:
    raise AssertionError("did not raise")

```

```python

```